# Module 1

#### In this assignment, you will work with ufo sightings data.
- The data includes various data points about individual ufo sightings
- Data File(s): ufo-sightings.csv

<!-- grade_id: init, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("ufo-sightings.csv")
from nose.tools import assert_equal
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
'''
1. Import the csv module. Load and read the UFO sightings data set, from the ufo-sightings.csv file, 
into a DictReader inside a with statement.  Assume the data file is in the same directory as the code. 

Print the field names of the data set. Iterate over the reader to put the data into a list name "ufosightings".

'''

import csv
filepath = "ufo-sightings.csv"

# your code here
with open(filepath, 'r') as csvfile:
    reader = csv.DictReader(csvfile)
    
    ufosightings = [] 
    for row in reader:
        ufosightings.append(row)
    
    print(ufosightings)
        
    
```

Output:
```
IOPub data rate exceeded.
The notebook server will temporarily stop sending output
to the client in order to avoid crashing it.
To change this limit, set the config variable
`--NotebookApp.iopub_data_rate_limit`.

Current values:
NotebookApp.iopub_data_rate_limit=1000000.0 (bytes/sec)
NotebookApp.rate_limit_window=3.0 (secs)
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings, sol.ufosightings)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
'''
2. How many sightings were there in total?  Put the count in "ufosightings_count" and print the result.
'''
# your code here

ufosightings_count = len(ufosightings)
print(ufosightings_count)
```

Output:
```
80332
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings_count, sol.ufosightings_count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
'''
3. How many sightings were there in the US?  Put the US sightings in "sightings_us" and print.

Hint: Check for ufo sightings where the country is 'us'.

'''

# your code here
sightings_us = [row for row in ufosightings if row["country"] == "us"]
print(len(sightings_us))
```

Output:
```
65114
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_us,  sol.sightings_us)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
'''
4. Let's find the "fireball" sighting(s) that lasted more than ten seconds in US. 
Print the the datetime and state of each.  Put the data in "fball" and print the result.

Note: Consider only the US sightings stored in "sightings_us".

- Cast the duration in seconds to a float (decimal). 
- Check if duration is greater than 10. 
- Check if the shape is "fireball".

'''

#First, define a Python function that checks if a given duration (seconds) is "valid"
def is_valid_duration(duration_as_string):
# your code here
    try:
        duration = float(duration_as_string)
    except ValueError:
        return False
    else:
        return duration

fball = [row for row in sightings_us if(is_valid_duration(row["duration (seconds)"]) > 10 and row["shape"] == "fireball")]

for event in fball:
    print(event["datetime"], event["state"]) 
```

Output:
```
10/10/1968 19:00 nc
10/10/1998 03:30 az
10/10/1998 23:50 tn
10/10/1999 11:00 ca
9/9/2010 22:00 nm
9/9/2011 20:00 ca
9/9/2012 19:05 tn
9/9/2012 20:00 nc
9/9/2013 20:20 al
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fball, sol.fball)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
'''
5. Sort the above list by duration. What was the datetime and duration of the longest sighting?  
Put the sorted list in "fballsorted" and print the result.

- Cast the duration in seconds to a float (decimal). 
- Sort in reverse order.

'''

# your code here
fballsorted = sorted(fball, key = lambda x: float(x["duration (seconds)"]), reverse=True)

# Print the sorted result
for entry in fballsorted:
    print(entry['datetime'], entry['duration (seconds)'])

```

Output:
```
8/3/2008 21:00 10526400
12/9/1965 24:00 6120000
7/5/1996 00:00 2631600
12/31/2013 21:54 172800
12/3/1999 19:30 145800
7/1/2001 02:48 37800
8/30/2003 21:00 37800
2/13/2012 18:00 21600
3/20/2000 22:00 21600
4/6/1978 20:00 21600
8/12/2007 03:00 21600
7/14/2000 23:30 18300
9/15/1968 24:00 18000
10/5/2013 18:00 14400
12/8/2005 13:00 14400
5/1/1994 20:00 14400
5/9/1981 00:00 14400
6/1/1978 18:00 14400
6/6/1960 12:00 14400
2/3/2012 18:00 600
3/10/2006 05:15 600
3/12/2013 22:50 600
3/16/2006 19:30 600
3/17/2012 23:30 600
3/18/2001 19:20 600
3/19/2011 21:00 600
3/19/2012 23:15 600
3/19/2014 00:40 600
3/20/2006 22:00 600
3/21/2012 22:00 600
3/23/2007 08:10 600
3/30/2003 21:00 600
3/3/2010 06:15 600
3/7/2004 01:37 600
3/7/2006 02:40 600
4/10/2011 22:00 600
4/1/2013 21:00 600
4/15/2011 20:00 600
4/16/2013 22:00 600
4/16/2014 10:30 600
4/17/2000 22:30 600
4/19/2004 21:10 600
4/20/2009 18:35 600
4/20/2010 21:00 600
4/20/2013 22:20 600
4/2/2012 18:20 600
4/22/2002 21:00 600
4/22/2014 20:00 600
4/23/2
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fballsorted, sol.fballsorted)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
'''
6. What state had the longest lasting "fireball"?   Put the state in "state" and print the result.

- Check if the shape is "fireball".
- Cast the duration in seconds to a float (decimal).
- Get the record with the largest (max) duration in seconds.
- Get the state for the record.

'''

# your code here
state = fballsorted[0]["state"]
print(state)
```

Output:
```
va
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(state, sol.state)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
'''
7. Let's assume that any sighting (of any shape) of 0 seconds is insignificant. 
Write code to filter out these extraneous records and get the shortest sighting overall now.  
Put the minimum duration in "min_duration" and print the result.  
Use ufosightings
Note: Consider all sightings stored in "ufosightings".

'''

# your code here

valid_sightings = [row for row in ufosightings if is_valid_duration(row["duration (seconds)"]) and float(row["duration (seconds)"]) > 0]

min_duration_entry = min(valid_sightings, key=lambda x: float(x["duration (seconds)"]))

min_duration = float(min_duration_entry["duration (seconds)"])

print(min_duration)

```

Output:
```
0.001
```

<!-- grade_id: q7_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(min_duration, sol.min_duration)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8, points: 0, grade: false, solution: true -->
```python
'''
8. What are the top 3 shapes sighted, and how many sightings were there for each? 

Note: Consider all sightings stored in "ufosightings".

- Create a new list "sightings_shapes" containing values from the "shape" column in ufosightings.  
- Create a new dictionary "count" with values of that column as keys and the counts as values.
- Get a list of the dictionary keys and values using the items() method.  This will return a list of key:value pairs.
Sort the list of key:value pairs in reverse order, from greatest (most sightings) to least.

Get the top 3 and store in "top3shapes".  Print the result.

'''

#Create a new list containing values from the "shape" column in ufosightings.
# your code here
sightings_shapes = [row["shape"] for row in ufosightings]


#Create a new dictionary with values of that column as keys and the counts as values.
# your code here
count = {}
for row in ufosightings:
    shape = row["shape"]
    count[shape] = count.get(shape, 0) + 1

        
#Get a list of the dictionary keys and values (use the items() method) and sort them in reverse order, from greatest (most sightings) to least.
#Get and print the top 3.
# your code here
sorted_count = sorted(count.items(), key=lambda x: x[1], reverse=True)

# Get and print the top 3
top3shapes = sorted_count[:3]
# for shape, count in top3shapes:
#     print(shape, count)
```

<!-- grade_id: q8_test_a, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_shapes, sol.sightings_shapes)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_b, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(count, sol.count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_c, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(top3shapes, sol.top3shapes)
print("Success!")
```

Output:
```
Success!
```